<a href="https://colab.research.google.com/github/toche7/BasicPython/blob/main/04_ControlFlow.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 04: Control Flow Statements

Version: 05 <br>
Date: May 9, 2024


## If



ลักษณะโครงสร้างของการใช้ if เป็นดังนี้

```
if some_condition:
    algorithm
```

ให้สังเกตว่าหลัง some\_condition จะตามด้วย colon \( : \)



```python
x = 12
if x >10:
    print("Hello")
```

<span style='color:#03a9f4'>เขียนโปรแกรมสั่งปริ้นชื่อ 5 ครั้ง</span><span style='color:#03a9f4'><span style='color:#3f51b5'> </span></span>


```python

```

```python

```

```python

```

```python

```

## If\-else



``` 
if some_condition:

    algorithm

else:

    algorithm
```


```python
x = 12
if x > 10:
    print("hello")
else:
    print("world")
```

<span style='color:#03a9f4'>เขียนโปรแกรมเพื่อตรวจสอบคุณสมบัติมีสิทธิเลือกตั้ง จากอายุ</span>


```python

```

```python

```

```python

```

## if\-elif



```
if some_condition:

    algorithm

elif some_condition:

    algorithm

else:

    algorithm
```


```python
x = 10
y = 12
if x > y:
    print("x>y")
elif x < y:
    print("x<y")
else:
    print("x=y")
```

ชุดคำสั่ง if ที่อยู่ภายใน ชุดคำสั่ง if\-elif หรือ if \-else จะถูกเรียกว่า "nested if statements"


```python
x = 10
y = 12
if x > y:
    print("x>y")
elif x < y:
    print("x<y")
    if x==10:
        print("x=10")
    else:
        print("invalid")
else:
    print ("x=y")
```

<span style='color:#03a9f4'>เขียนโปรแกรม เพื่อขัดแยกระดับคุณภาพอาการเป็น อันตราย พอใช้ ดี จากค่า ​Pm2.5</span>


```python

```

```python

```

```python

```

```python

```

## Loops



### For



``` 
for variable in something:
    
    algorithm

```python
for i in range(5):
    print(i)
```

<span style='color:#03a9f4'>จงเขียนโปรแกรมเพื่อบวกเลข 1\-10</span>


```python

```

```python

```

```python

```

```python

```

### While



วิธีการการใช้งาน while loop เป็นดังนี้

```
while some_condition:

    algorithm



```python
i = 1
while i < 3:
    print(i ** 2)
    i = i+1
print('Bye')
```

<span style='color:#03a9f4'>จงเขียนโปรเพื่อตรวจสอบคุณภาพอากาศหาก เป็นระดับ ระวัง ให้ปริ้นข้อความเตือนออกมา หากอยู่ในระดับดีไม่ต้องปริ้นอะไรออกมา </span>

use function  "input\( \)" to get the input value.


```python

```

```python

```

```python

```

```python

```

## Break



คำสั่ง Break จะใช้ในการหยุดการทำงานที่อยู่ภายใน loop และจะหลุดออกมาจาก loop 



```python
for i in range(100):
    print (i)
    if i>=7:
        break
```

<span style='color:#03a9f4'>จงเขียนโปรเพื่อตรวจสอบการรับค่า keyboard  จำนวน 10 รอบ หากมีการรับค่า key  = 'y' ให้ทำหลุดออกมาจากลูป</span> 


```python

```

```python

```

```python

```

## Continue

คำสั่ง continue จะทำให้เกิดการข้ามคำสั่งส่วนที่เหลือภายใน loop นั้นกลับไปทำงานในรอบใหม่  โดยคำสั่งนี้จะใช้ได้ทั้ง while และ for เช่นเดียวกับคำสั่ง break


```python
for i in range(10):
    if i>4:
        print("Ignored",i)
        continue
    # this statement is not reach if i > 4
    print("Processed",i)
```

<span style='color:#03a9f4'>จงเขียนโปรเพื่อตรวจสอบการรับค่า keyboard  จำนวน 3 รอบ หากมีการรับค่า password  = 'mypass' ถูกต้องจะแสดงค่าว่า  "start the system", หากยังไม่ถูกให้กลับไปรับถ้าค่าใหม่ แต่ถ้าเกิน  3 ครั้งให้แสดงคำว่า "Your account is locked, bye bye"</span>


```python

```

```python

```

```python

```

## Exercise



```python
# for colab please uncomment the following command
#!pip install nose
```

## Ex 1

จงเขียนโปรแกรมเพื่อตัดเกรดให้กับวิชา Basic Programing  โดยมีเงื่อนไขว่า ถ้าได้คะแนนตั้ง 80  ขึ้นไปจะได้เกรด A ถ้าต่ำกว่า 80 จะได้เกรด B


<!-- grade_id: aa4de8, points: 0, grade: false, solution: true -->
```python
def exercise1(score = 70):
    '''
    score : raw score 
    grade : grade that compute form score 
    score_A: mine score of Grade A
    '''
    score_A = 80
    grade   = 'B'
    ### Your code here
    

    return grade, score
    
```

<!-- grade_id: 073cc7, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(exercise1(), ('B',70))
assert_equal(exercise1(80), ('A',80))
print("Grade %s for score %d" %exercise1(80))
```

Output:
```
Grade A for score 80
```

### Ex 2

จาก Ex1 หากต้องการให้เกรดครบทั้ง A B C D โดยมีระดับคะแนนสำหรับเกรดคือ 80 70 60 50 ตามลำดับ หากได้คะแนนต่ำกว่า 50 จะได้เกรด E จงเขียนโปรแกรมทีสามารถให้ครบทั้ง 5 ระดับ


<!-- grade_id: 413878, points: 0, grade: false, solution: true -->
```python
def grading(score = 70):
    '''
    score : raw score 
    grade : grade that compute form score 
    score_A: mine score of Grade A
    '''
    score_A = 80
    score_B = 70
    score_C = 60
    score_D = 50
    grade   = 'E'
    ## Your Code Here

    return grade, score
```

```python
grading(55)
```

<!-- grade_id: efb8c5, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(grading(40), ('E',40))
assert_equal(grading(55), ('D',55))
assert_equal(grading(65), ('C',65))
assert_equal(grading(75), ('B',75))
assert_equal(grading(85), ('A',85))
print("Grade %s for score %d" %grading(85))
```

Output:
```
Grade A for score 85
```

```python

```