<a href="https://colab.research.google.com/github/toche7/BasicPython/blob/main/05_DataStructureP1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 05: Data Struct Part 1



Version: 06 <br>


## Lists

Lists เป็นโครงสร้างข้อมูลที่ใช้บ่อยที่สุด คิดว่าเป็นลำดับของข้อมูลที่อยู่ในวงเล็บเหลี่ยมและข้อมูลจะถูกคั่นด้วยลูกน้ำ แต่ละข้อมูลเหล่านี้สามารถเข้าถึงได้โดยเรียกว่าค่าดัชนี

list จะถูกประกาศด้วยเครื่องหมาย '[  ]' อย่างเช่น

```python
a = []
```

```python
print(type(a))
```

```python

```

```python

```

```python

```

คำสั่ง type\( \)  จะใช้ตรวจสอบชนิดหรือ class ของ object ที่ต้องการตรวจสอบ นอกจากนั้นแล้วเรายังใช้คำสั่ง  dir\( \) ในการตรวจเช็คว่า object ดังกล่าวมี attribute หรือ method อะไรบ้างดังนี้


```python
dir(a)
```

```python

```

```python

```

```python

```

```python

```

### Indexing

กลับมาที่ Lists เราสามารถกำหนดลำดับของข้อมูลให้กับ list x ได้โดยตรงดังที่แสดง


```python
x = ['apple', 'orange']
```

ใน python การจัดทำดัชนีจะเริ่มจาก 0 ตอนนี้ List x ซึ่งมีสองสมาชิกจะมีแอปเปิ้ลที่ 0 ดัชนีและสีส้มที่ 1 ดัชนี



```python
x[0]
```

```python

```

```python

```

```python

```

การจัดทำดัชนีสามารถทำได้ในลำดับย้อนกลับ นั่นคือสมาชิกสุดท้ายสามารถเข้าถึงได้ก่อน ที่นี่การสร้างดัชนีเริ่มจาก -1 ดังนั้นค่าดัชนี -1 จะเป็นสีส้มและดัชนี -2 จะเป็นแอปเปิ้ล



```python
x[-1]
```

```python

```

```python

```

```python

```

อย่างที่เราอาจเดาได้แล้ว x [0] = x [-2], x [1] = x [-1] แนวคิดนี้สามารถขยายไปยัง List ที่มีสมาชิกอื่นๆ อีกมาก

```python
y = ['carrot','potato']
```

```python

```

```python

```

```python

```

ตัวอย่างนี่เราได้ประกาศ List x และ y แต่ละ List มีข้อมูลของตัวเอง ตอนนี้สอง List นี้สามารถใส่ใน List อื่นได้ เช่น z ซึ่งจะมีข้อมูลเป็นสอง list โดยเราจะเห็นว่านี้ภายใน list เรียกว่าเป็น list ที่ซ้อนกันและเป็นวิธีการประกาศอาร์เรย์ซึ่งเราจะเห็นในภายหลัง



```python
z  = [x,y]
print (z)
```

```python

```

```python

```

```python

```

การทำดัชนีใน list ที่ซ้อนกันอาจสร้างความสับสนได้หากเราไม่เข้าใจว่าการจัดทำดัชนีทำงานอย่างไรใน python ซึ่งอยากให้ทำความเข้าใจจากตัวอย่างดังต่อไปนี้

ให้เราเข้าถึงข้อมูล 'apple' ใน list ที่ซ้อนด้านบน
อันดับแรก ที่ดัชนี 0 มี list \['apple','orange'\] และที่ดัชนี 1 มี list อื่น \['carrot','potato'\] ดังนั้น z\[0\] ควรให้ list แรกที่มี 'apple' แก่เรา



```python
z1 = z[0]
print (z1)
```

ตอนนี้ให้สังเกตว่า z1 ไม่ใช่ list ที่ซ้อนกัน ดังนั้นเพื่อเข้าถึง 'apple' z1 ควรจะสร้างดัชนีที่ 0



```python
z1[0]
```

```python

```

```python

```

```python

```

แทนที่จะทำด้านบน ใน python เราสามารถเข้าถึง 'apple' ได้โดยเพียงแค่เขียนค่าดัชนีในแต่ละครั้งเคียงข้างกัน



```python
z[0][0]
```

```python

```

```python

```

```python

```

หากมี list อยู่ใน list ภายใน list เราสามารถเข้าถึงค่าในสุดได้โดยดำเนินการ

```
z[ ][ ][ ]
```



### Slicing



การจัดทำดัชนีจำกัดการเข้าถึงเพียงค่าเดียว ในทางกลับกัน Slicing สามารถกำลังเข้าถึงข้อมูลหลายค่าได้ ซึ่งเรียกว่า "การหั่น"  list  (Slicing the list)

Slicing ทำได้โดยกำหนดค่าดัชนีของสมาชิกแรกและสมาชิกสุดท้ายจาก list หลักที่จำเป็นใน list แบบแบ่งส่วน มันถูกเขียนเป็น parentlist[ a : b ] โดยที่ a,b คือค่าดัชนีจาก parentlist หากไม่ได้กำหนด a หรือ b ค่าดัชนีจะถือเป็นค่าแรกสำหรับ a หากไม่ได้กำหนด a และค่าสุดท้ายสำหรับ b เมื่อไม่ได้กำหนด b



```python
num = [0,1,2,3,4,5,6,7,8,9]
```

```python
print (num[0:4])
print (num[4:])
```

เรายังสามารถแบ่ง list หลัก โดยกำหนดความของแต่ละ step ได้



```python
num[:9:3]
```

```python

```

```python

```

```python

```

### Built in List Functions

หากต้องการค้นหาความยาวของ list หรือจำนวนสมาชิกใน list ให้ใช้ **len( )**

```python
len(num)
```

```python

```

```python

```

```python

```

หาก list ประกอบด้วยสมาชิกประเภทจำนวนเต็มทั้งหมด ดังนั้น **min( )** และ **max( )** จะให้ค่าต่ำสุดและสูงสุดใน list

```python
min(num)
```

```python
max(num)
```

```python

```

```python

```

```python

```

list สามารถต่อกันได้โดยการเพิ่ม '+' ผลลัพธ์จะมีสมาชิกทั้งหมดของ list ที่เพิ่มเข้ามา  list ผลลัพธ์จะไม่เป็น list ที่ซ้อนกัน

```python
[1,2,3] + [5,4,7]
```

```python

```

```python

```

```python

```

บางครั้ง อาจมีข้อกำหนดที่เราต้องตรวจสอบว่ามีสมาชิกเฉพาะใน list ที่กำหนดไว้ล่วงหน้าหรือไม่ พิจารณา list ด้านล่างนี้


```python
names = ['Earth','Air','Fire','Water']
```

เพื่อตรวจสอบว่า 'Fire' และ 'Air' มีอยู่ใน list หรือไม่ วิธีการทั่วไปคือการใช้ for loop และวนซ้ำ list และใช้เงื่อนไข if แต่ใน python เราสามารถใช้แนวคิด 'a in b' ซึ่งจะคืนค่า 'True' หาก a มีอยู่ใน b และ 'False' หากไม่มี


```python
'Fire' in names
```

```python

```

```python

```

```python

```

ใน list ที่มีสมาชิกเป็นสตริง **max( )** และ **min( )** สามารถใช้ได้ **max( )** จะส่งคืนสมาชิกสตริงที่มีค่า ASCII สูงสุดและต่ำสุดเมื่อใช้ **min( )** โปรดสังเกตุว่าจะมีการพิจารณาเฉพาะดัชนีแรกของแต่ละสมาชิกในแต่ละครั้ง และหากค่าของดัชนีเท่ากัน ดัชนีที่สองจะพิจารณาเป็นเช่นนี้ไปเรื่อยๆ

```python
mlist = ['bzaa','ds','nc','az','z','klm']
```

```python
print (max(mlist))
print (min(mlist))
```

```python

```

```python

```

```python

```

ที่นี่พิจารณาดัชนีแรกของแต่ละสมาชิกและด้วยเหตุนี้ z มีค่า ASCII สูงสุดจึงถูกส่งกลับและ ASCII ขั้นต่ำคือ a แต่ถ้าตัวเลขถูกประกาศเป็นสตริงล่ะ?

```python
nlist = ['1','94','93','1000']
```

```python
print (max(nlist))
print (min(nlist))
```

แม้ว่าตัวเลขจะถูกประกาศเป็นสตริง ดัชนีแรกของแต่ละสมาชิกจะถูกพิจารณาและส่งคืนค่าสูงสุดและต่ำสุดตามลำดับ

```python

```

```python

```

```python

```

แต่ถ้าเราต้องการค้นหาสมาชิกแบบสตริง **max( )** ตามความยาวของสตริง พารามิเตอร์ 'key=len' จะถูกประกาศภายใน **max( )** และ **min( )** เพื่อหาค่าดังกล่าว

```python
print (max(names, key=len))
print (min(names, key=len))
```

แต่แม้ว่า  Water ก็มีความยาว 5  **max\(\)** หรือ **min\(\)** ฟังก์ชันจะคืนค่าสมาชิกแรกเมื่อมีสมาชิกตั้งแต่สองสมาชิกขึ้นไปที่มีความยาวเท่ากัน <br>
สามารถใช้ฟังก์ชันอื่นๆ ในตัวหรือฟังก์ชัน lambda \(จะกล่าวถึงในภายหลัง\) แทน len

```python

```

```python

```

```python

```

สตริงสามารถแปลงเป็น list โดยใช้ฟังก์ชัน **list\(\)**

```python
list('hello')
```

```python

```

```python

```

```python

```

**สรุปฟังก์ชันที่สามารถใช้กับ list ได้**

- min\(lst\) : คืนค่าที่น้อยที่สุดออกมา
- max\(lst\) :  คืนค่าที่มากที่สุดออกมา
- len\(lst\) : คืนค่าความยาว lst
- list\(seq\) : แปลงค่าข้อมูลประเภท seq ให้ไปเป็น list เช่นข้อมูลประเภท  range, string และ tuple



### Method ของ List



method เป็นฟังก์ชันที่สร้างขึ้นมาเพื่อใช้งานกับ object ประเภท list โดยเฉพาะ เพื่อไม่ให้เกิดความสับสนเราจะเรียนฟังก์ชันมีมาให้กับ object นั้นว่า method แทนที่จะเรียกว่าฟังก์ชัน เพราะวิธีการเรียกใช้งานต่างกัน เช่น

`a = [ ]`

`a.append('hello')`

จากตัวอย่างทางด้านล่างนี้เราจะเห็นว่า append() เป็น method ของ object a จะต้องเรียกใช้ด้วยวิธีการนี้



```python

```

```python

```

```python

```

**append\( \)** ใช้เพื่อเพิ่มสมาชิกที่ส่วนท้ายของ list



```python
lst = [1,1,4,8,7]
```

```python
lst.append(1)
print(lst)
```

```python

```

```python

```

```python

```

**count\( \)** ใช้เพื่อนับจำนวนของสมาชิกเฉพาะที่มีอยู่ใน list


```python
lst.count(1)
```

```python

```

```python

```

```python

```

ฟังก์ชัน **append\( \)** สามารถใช้เพื่อเพิ่ม list ทั้งหมดในตอนท้ายได้ สังเกตว่า list ผลลัพธ์ กลายเป็น list ที่ซ้อนกัน


```python
lst1 = [5,4,2,8]
```

```python
lst.append(lst1)
print(lst)
```

```python

```

```python

```

```python

```

แต่ถ้า list ที่ซ้อนกัน ไม่ใช่สิ่งที่ต้องการ ก็สามารถใช้ฟังก์ชัน **extend\( \)** ได้


```python
lst.extend(lst1)
print(lst)
```

```python

```

```python

```

```python

```

**index\( \)** ใช้เพื่อค้นหาค่าดัชนีของสมาชิกเฉพาะ โปรดทราบว่าหากมีสมาชิกหลายสมาชิกที่มีค่าเดียวกัน ค่าดัชนีแรกของสมาชิกนั้นจะถูกส่งคืน


```python
lst.index(1)
```

```python

```

```python

```

```python

```

**insert\(x,y\)** ใช้เพื่อแทรกสมาชิก y ที่ค่าดัชนีที่ระบุ x สำหรับฟังก์ชัน **append\( \)** ทำให้แทรกได้เฉพาะส่วนท้ายเท่านั้น


```python
lst.insert(5, 'name')
print(lst)
```

```python

```

```python

```

```python

```

**insert\(x, y\)** แทรก แต่ไม่ได้แทนที่สมาชิก หากเราต้องการแทนที่สมาชิกด้วยสมาชิกอื่น เราเพียงแค่กำหนดค่าให้กับดัชนีนั้น ๆ


```python
lst[5] = 'Python'
print(lst)
```

```python

```

```python

```

```python

```

ฟังก์ชัน **pop\( \)** คืนค่าสมาชิกสุดท้ายใน list ซึ่งคล้ายกับการทำงานของ stack ดังนั้นจึงไม่ผิดที่จะบอกว่า list สามารถใช้เป็น stack ได้


```python
lst.pop()
```

```python

```

```python

```

```python

```

เราสามารถระบุค่าดัชนีเพื่อแสดงสมาชิกที่สอดคล้องกับค่าดัชนีนั้น


```python
lst.pop(0)
```

**pop\( \)** ใช้เพื่อลบสมาชิกตามค่าดัชนีซึ่งสามารถกำหนดให้กับตัวแปรได้ นอกจากนี้ยังสามารถลบสมาชิกโดยการระบุสมาชิกโดยใช้ฟังก์ชัน **remove\(\)**



```python
lst.remove('Python')
print(lst)
```

```python

```

```python

```

```python

```

นอกจากนั้น เรายังมีวิธีการอื่นเพื่อลบ คือ **del** แต่ del นี้ไม่ได้เป็น method  ของ list


```python
del lst[1]
print (lst)
```

```python

```

```python

```

```python

```

สมาชิกทั้งหมดที่มีอยู่ใน list สามารถย้อนกลับได้โดยใช้ฟังก์ชัน **reverse\(\)**


```python
lst.reverse()
print(lst)
```

```python

```

```python

```

```python

```

สังเกตุว่า nested list  \[5,4,2,8\] จะถือเป็นสมาชิกเดียวของ list หลัก lst ดังนั้นสมาชิกภายใน nested list จะไม่กลับ list

Python เสนอการดำเนินการ **sort\( \)** ในตัวเพื่อจัดเรียงสมาชิกในลำดับจากน้อยไปมาก



```python
lst.sort()
print(lst)
```

อย่างไรก็ตาม **sort\(\)** ไม่รองรับการเปรียบเทียบข้อมูลต่างประเภทกัน ดังนั้นจะต้องเอาส่วนที่เป็น nested list ออกไปก่อน



```python
lst.remove([5, 4, 2, 8])
print(lst)
```

```python
lst.sort()
print(lst)
```

สำหรับลำดับจากน้อยไปมาก มีค่าเริ่มต้น เงื่อนไขย้อนกลับจะเป็น False ดังนั้นการเปลี่ยนเป็น True จะจัดเรียงสมาชิกจากมากไปหาน้อย


```python
lst.sort(reverse=True)
print(lst)
```

```python

```

```python

```

```python

```

ในทำนองเดียวกันสำหรับ list ที่มีสมาชิกสตริง **sort\( \)** จะเรียงลำดับสมาชิกตามค่า ASCII จากน้อยไปมาก และโดยการระบุ reverse=True เพื่อจัดเรียงจากมากไปหาน้อย



```python
names.sort()
print(names)

```

```python
names.sort(reverse=True)
print(names)
```

```python

```

```python

```

```python

```

ในการจัดเรียงตามความยาว key=len ควรระบุตามที่แสดง



```python
names.sort(key=len)
print(names)
```

```python
names.sort(key=len,reverse=True)
print(names)
```

```python

```

```python

```

```python

```

**สรุป method ของ list**

- append\( \) :
- count\( \)
- extend\( \)
- index\( \)
- insert\( \)
- pop\( \)
- remove\( \)
- reverse\( \)
- sort\( \)



```python

```

### Copying a list



โปรแกรมเมอร์ python ใหม่ ส่วนใหญ่จะทำข้อผิดพลาดนี้ พิจารณาสิ่งต่อไปนี้



```python
lista= [2,1,4,3]
```

```python
listb = lista
print (listb)
```

ตรงนี้ เราได้ประกาศ lista = \[2,1,4,3\]  list นี้ถูกคัดลอกไปยัง listb โดยกำหนดค่าและได้รับการคัดลอกตามที่เห็น ตอนนี้เราทำการเปลี่ยนค่าบน lista


```python
lista.pop()
print (lista)
lista.append(9)
print (lista)
```

```python
print(listb)
```

listb ก็เปลี่ยนไปเช่นกันแม้ว่าจะไม่ได้ดำเนินการใดๆ เนื่องจากเราได้กำหนดพื้นที่หน่วยความจำเดียวกันกับ lista ใน listb แล้วจะแก้ไขอย่างไร?

หากยังจำได้ ในการสไลซ์ เราพบว่า parentlist\[a:b\] ส่งคืน list จาก parent list ทีมีดัชนีเริ่มต้น a และดัชนีสิ้นสุด b และหากไม่มีการระบุ a และ b โดยค่า default จะถือว่าเป็นสมาชิกแรกและสมาชิกสุดท้าย เราใช้แนวคิดเดียวกันที่นี่ ดังนั้นเราจะกำหนดข้อมูลของ lista ให้กับ listb เป็นตัวแปร



```python
lista = [2,1,4,3]
```

```python
listb = lista[:]
print (listb)
```

```python
lista.pop()
print (lista)
lista.append(9)
print (lista)
```

```python
print (listb)
```

## Exercise



### Ex 1

ถ้าห้องเรียน ป. 1/1  มีรายชื่อนักเรียนดังนี้

สมชาย สมหญิง สมศรี  

ถ้าห้องเรียน  ป. 1/2 มีรายชื่อนักเรียนดังนี้

สมศักดิ์ สมพงษ์ สมรักษ์

ต้องการสร้าง list รายชื่อของ ห้องเรียน ป. 1/1 และ ป. 1/2 และรายชื่อรวมของชั้น ป.1 ทั้งหมด



<!-- grade_id: 82fde7, points: 0, grade: false, solution: true -->
```python
def exercise1():

    nameListClass1_1 = [ ]
    nameListClass1_2 = [ ]
    allNameList = []
    #===================================================
    ## Your code starts here

    #===================================================

    return nameListClass1_1, nameListClass1_2, allNameList
```

<!-- grade_id: 894cf5, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:

assert exercise1()[0] == ['สมชาย','สมหญิง','สมศรี']
print(exercise1())

### BEGIN HIDDEN TESTS
# students will NOT see these extra tests
assert exercise1()[1] == ['สมศักดิ์','สมพงษ์','สมรักษ์']
assert exercise1()[2] == ['สมชาย','สมหญิง','สมศรี','สมศักดิ์','สมพงษ์','สมรักษ์']
### END HIDDEN TESTS
```

### Ex 2

ร้านค้าแห่งหนึ่งมียอดขายประจำปี 2564 ในหน่วยล้านบาท ดังนี้

1.00 2.00 3.00 2.00 1.20 1.35 2.50 3.11 2.00 2.15 1.11 1.42

จงเก็บขอมูลต่อไปนี้ไว้ในตัวแปร list <br>
เมื่อมีข้อมูลในเดือนถัดไปเพิ่มเข้ามาให้นำเข้าไปต่อใน list นี้ อย่างเช่น 1.30 <br>
ให้แสดงผลข้อมูลจากน้อยไปมาก หลังจากที่มีการเพิ่มข้อมูลในเดือนใหม่แล้ว



<!-- grade_id: 57c50f, points: 0, grade: false, solution: true -->
```python
def exercise2():
    """
    Compute the squares of the numbers from 1 to n.  [replace with function description]
    """
    lstIncome = [1.00,2.00,3.00,2.00,1.20, 1.35, 2.50, 3.11, 2.00, 2.15, 1.11, 1.42]
    newIncome = 1.30
    #===================================================
    ## Your code starts here

    #===================================================
    return lstIncome
```

<!-- grade_id: ea064d, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:

assert exercise2()[12] ==  3.11
print(exercise2())

### BEGIN HIDDEN TESTS
# students will NOT see these extra tests
assert exercise2() == [1.0, 1.11, 1.2, 1.3, 1.35, 1.42, 2.0, 2.0, 2.0, 2.15, 2.5, 3.0, 3.11]
### END HIDDEN TESTS
```

Output:
```
[1.0, 1.11, 1.2, 1.3, 1.35, 1.42, 2.0, 2.0, 2.0, 2.15, 2.5, 3.0, 3.11]
```