<a href="https://colab.research.google.com/github/toche7/BasicPython/blob/main/06_DataStructureP2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 06: Data Structures Part 2



Version: 05 <br>
Date: May 9, 2024


## Tuples



tuple นั้นคล้ายกับ list แต่ความแตกต่างที่สำคัญเพียงอย่างเดียวคือสมาชิกภายใน list สามารถเปลี่ยนแปลงได้ แต่ใน tuple จะไม่สามารถเปลี่ยนแปลงได้ ให้คิดว่า tuples เป็นสิ่งที่ต้องเป็นจริงสำหรับบางสิ่งบางอย่าง และไม่สามารถจริงได้สำหรับค่าอื่น ๆ  เพื่อความเข้าใจที่ดีขึ้น ให้ทบทวนฟังก์ชัน   **divmod()**



```python
xyz = divmod(10,3)
print(xyz)
```

```python
print(type(xyz))
```

ในที่นี้ผลหารต้องเป็น 3 และส่วนที่เหลือต้องเป็น 1 ค่าเหล่านี้ไม่สามารถเปลี่ยนแปลงได้เมื่อ 10 ถูกหารด้วย 3 ดังนั้น divmod จะส่งกลับค่าเหล่านี้ใน tuple

ในการกำหนด tuple ตัวแปรถูกกำหนดด้วยเครื่องหมายวงเล็บ ( ) หรือ tuple( )

```python
tup = ()
```

```python
tup2 = tuple()
```

```python

```

```python

```

```python

```

```python

```

หากเราต้องการประกาศ tuple โดยตรง สามารถทำได้โดยใช้เครื่องหมายจุลภาคที่ส่วนท้ายของข้อมูล

```python
27,
```

```python
tt = (27,)
```

```python
type(tt)
```

> เราจะเขียน tt = \(27,\)  เพื่อแทน tuple ที่มีสมาชิก 1 ตัว ให้สังเกตุว่ามีการใส่เครื่องหมาย comma 
>
> แต่เราจะไม่เขียนเป็น tt = \(27\) เนื่องจากระบบจะมองเป็นวงเล็บในการคำนวณทั่วไปไม่ได้หมายถึง tuple



```python
tup3 = (1) ## ต้องระวัง
```

```python
a = 1,2,3
```

```python

```

```python

```

```python

```

```python

```

### Tuple จะแก้ไขค่าไม่ได้

โดยเราจะลองเปรียบเทียบกับ list ดังนี้


```python
lst = [10,20]
```

```python
tup = (1,2)
```

```python

```

```python

```

```python

```

### การคูณและบวก Tuple

27 เมื่อคูณด้วย 2 ได้ 54 แต่เมื่อคูณด้วย tuple ข้อมูลจะถูกทำซ้ำสองครั้ง



```python
2*(27,)
```

```python

```

```python
27,+ 28, + 29,
```

```python

```

```python

```

### การแปลงชนิดข้อมูลไปเป็น Tuple



สามารถกำหนดค่าได้ในขณะที่ประกาศ tuple ใช้ list เป็นอินพุตและแปลงเป็น tuple หรือนำสตริงมาแปลงเป็น tuple

```python
tup3 = tuple([1,2,3])
print(tup3)
```

```python
tup4 = tuple('Hello')
print(tup4)
```

**create tub5, tup6**

```python

```

```python

```

```python

```

### Index และ Slicing สำหรับ Tuple



tuple ใช้หลักการเดี่ยวกับ list


```python
print(tup3[1])
```

```python
tup5 = tup4[:3]
```

```python
print(tup5)
```

**create tub5, tup6**


```python

```

```python

```

```python

```

```python

```

### Mapping one tuple to another



```python
(a,b,c)= ('alpha','beta','gamma')
```

```python
print(a,b,c)
```

```python

```

```python

```

```python

```

```python
d = tuple('RajathKumarMP')
print(d)
```

```python

```

```python

```

```python

```

### Method ของ Tuple



ฟังก์ชัน **count\(\)** จะนับจำนวนสมาชิกที่ระบุซึ่งมีอยู่ใน tuple



```python
d.count('a')
```

```python

```

```python

```

```python

```

ฟังก์ชัน **index\(\)** คืนค่าดัชนีของสมาชิกที่ระบุ หากสมาชิกมากกว่าหนึ่ง list  ดัชนีของสมาชิกแรกของสมาชิกที่ระบุจะถูกส่งกลับ


```python
d.index('a')
```

```python

```

```python

```

```python

```

### Built in Tuple Functions



```python
dir(d)
```

```python
min(d)
```

```python
max(d)
```

```python
len(d)
```

```python

```

```python

```

```python

```

## Sets



การประกาศข้อมูลแบบ set


```python
setA = {1}
setB = set()
setC = {} ## อันนี้ต้องระวัง
```

```python
print(type(setA))
print(type(setB))
print(type(setC))
```

```python
set1 = set()
print(type(set1))
```

```python

```

```python

```

```python

```

set ส่วนใหญ่จะใช้เพื่อกำจัดตัวเลขซ้ำใน sequence หรือ list นอกจากนี้ยังใช้เพื่อดำเนินคำสั่ง set บางอย่าง

set ถูกประกาศเป็น set\(\) ซึ่งจะเริ่มต้น set ว่าง นอกจากนี้ยังสามารถดำเนินการ set\(\[sequence\]\) เพื่อประกาศ set ที่มีสมาชิก



```python
set0 = set([1,2,2,3,3,4])
print(set0)
```

สมาชิก 2,3 ที่ทำซ้ำสองครั้งจะเห็นเพียงครั้งเดียว ดังนั้นใน set แต่ละสมาชิกจะแตกต่างกัน


```python
st1 = {1}
type(st1)
```

```python

```

```python

```

```python

```

### Method ของ Set



```python
set1 = set([1,2,3])
```

```python
set2 = set([2,3,4,5])
```

**create set4, set5**

```python

```

```python

```

```python

```

```python

```

```python
dir(set1)
```

ฟังก์ชัน **union\( \)** ส่งคืน set ที่มีสมาชิกทั้งหมดของทั้งสอง set โดยไม่มีการทำซ้ำ


```python
set1.union(set2)
```

```python

```

```python

```

```python

```

**add\( \)** จะเพิ่มสมาชิกเฉพาะเข้าไปใน set สังเกตุว่าดัชนีของสมาชิกที่เพิ่มใหม่นั้นจะเกิดในตำแหน่งต่างๆ ไม่จำเป็นในท้ายที่สุด


```python
set1.add(0)
set1
```

```python

```

```python

```

```python

```

ฟังก์ชัน **intersection\( \)** คืนค่า set ที่มีสมาชิกทั้งหมดที่อยู่ในทั้ง set ชุด


```python
set1.intersection(set2)
```

```python

```

```python

```

```python

```

ฟังก์ชัน **difference\( \)** ดึง set ที่มีสมาชิกที่อยู่ใน set1 และไม่อยู่ใน set2


```python
set1.difference(set2)
```

```python
set2.difference(set1)
```

```python

```

```python

```

```python

```

ฟังก์ชัน **symmetric\_difference\( \)** แสดงผลฟังก์ชันที่มีสมาชิกที่อยู่ในชุดใดชุดหนึ่ง


```python
set2.symmetric_difference(set1)
```

```python

```

```python

```

```python

```

**issubset\( \), isdisjoint\( \), issuperset\( \)** ใช้เพื่อตรวจสอบว่า set1/set2 เป็น subset, disjoint หรือ superset ของ set2/set1 ตามลำดับ


```python
set1.issubset(set2)
```

```python
set2.isdisjoint(set1)
```

```python
set2.issuperset(set1)
```

```python

```

```python

```

```python

```

pop\( \) ใช้เพื่อลบสมาชิกที่กำหนดเองใน set


```python
set1.pop()
print(set1)
```

```python

```

```python

```

```python

```

ฟังก์ชัน **remove\( \)** ลบสมาชิกที่ระบุออกจาก set


```python
set1.remove(2)
set1
```

```python

```

```python

```

```python

```

**clear\( \)** ใช้เพื่อล้างสมาชิกทั้งหมดและทำให้ set นั้นเป็น set ว่าง



```python
set1.clear()
set1
```

```python

```

```python

```

```python

```

## Exercise



### Ex 1

หากมีการประกาศค่าตัวแปรแบบ tuple ดังนี้ <br>
tuple1: (99, 88)<br>
tuple2: (11, 22)<br>

และต้องการสลับค่าระหว่าง 2 ตัวแปรดังนี้ <br>
tuple1: (11, 22)<br>
tuple2: (99, 88)<br>

จงเขียนโปรแกรมเพื่อสลับค่าตัวแปรดังกล่าวนี้



<!-- grade_id: 5d1d52, points: 0, grade: false, solution: true -->
```python
def exercise3():
    """
    """
    tp1 = (99, 88)
    tp2 = (11, 22)
    #===================================================
    ### Your code here
    
    #===================================================
    return tp1,tp2
```

<!-- grade_id: 7649b3, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(exercise3(), ((11, 22),(99, 88)))


```

### Ex 2

โรงเรียนแห่งหนึ่ง เปิดให้มีกิจกรรมชมรมขึ้น จำนวน 2 ชมรมคือ X และ Y โดยนักเรียนสามารถที่จะเข้าไปสมาชิกได้ 2 ชมรม โดยที่ชมรม X มีจำนวนสมาชิกประกอบด้วย a b c d

และชมรม Y มีจำนวนสมาชิกประกอบด้วย  a c e f g

จงเขียนโปรแกรมเพื่อแสดงเฉพาะสมาชิกที่อยู่ในแต่ละชมรม และที่อยู่ทั้ง 2 ชมรมนี้ 


<!-- grade_id: 074dce, points: 0, grade: false, solution: true -->
```python
def exercise5():  # modify function name and parameters
    """
    clubX : member of club X
    clubY : memeber of club Y
    itsMember : intersect member on both club
    """
    clubX = {}
    clubY = {}
    itsMember = {}
    #===================================================
    ### Your code here
    
    #===================================================
    return clubX, clubY, itsMember
```

<!-- grade_id: db77b5, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(exercise5(),({'a', 'b', 'c', 'd'}, {'a', 'c', 'e', 'f', 'g'}, {'a', 'c'}))

```

```python

```

```python

```