<a href="https://colab.research.google.com/github/toche7/BasicPython/blob/main/07_DataStructureP3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 07: Data Structures Part 3



Version: 06 <br>

## Dictionaries



Dictionaries มีการใช้งานเหมือนกับ database เพราะเราสามารถเข้าถึงข้อมูลด้วย สตริง \(string\) ที่ได้กำหนดไว้ล่วงหน้าแล้ว



การกำหนด ค่าใน dictionary สามารถทำได้ยใช้เครื่องหมาย { } หรือคำสั่ง dict\(\)


```python
d0 = {}
d1 = dict()
print(type(d0), type(d1))
```

```python

```

```python

```

d2 = {'a': 1, 'b': 2, 'c': 3}

```python

```

```python

```

```python

```

d3 = dict(a=1, b=2, c=3)

```python

```

```python

```

Dictionary จะมีการทำงานเหมือนกับการใช้ list แต่จะเพิ่มความสามารถด้วยการกำหนดลักษณะของ index ได้



```python
d0
```

```python
d0['One'] = 1
d0['OneTwo'] = 12
print(d0)
```

```python

```

```python

```

```python

```

จากตัวอย่างข้างคนเป็นการกำหนดคำ index ด้วยคำ word และมีค่าเป็นตัวเลข  การเรียกใช้ค่า 1 สามารถเรียกด้วยการระบุเป็นสตริง 'One'


```python
print (d0['One'])
```

```python

```

```python

```

```python

```

เราสามารถใช้ list สองอันมารวมกันเพื่อสร้างเป็น Dictionary ได้


```python
names = ['One', 'Two', 'Three', 'Four', 'Five']
numbers = [1, 2, 3, 4, 5]
```

```python

```

```python

```

```python

```

**zip\( \)**  ฟังก์ชั้นนี้สามารถสร้าง Dictionary จาก list สองตัวได้



```python
z1 = zip(names,numbers)
print(z1)
```

```python

```

List ทั้ง สองตัว ถูกรวมเข้าด้วยกันเพื่อสร้าง List เดียว และแต่ละ สมาชิกภายใน list จะถูกรวมเข้ากับ สมาชิกจาก list อื่นๆ ในตำแหน่งเดียวเป็นไปเป็น tuple เนื่องจาก tuple จะไม่สามารถเปลี่ยนค่าได้

การนำไปใช้จะต้องมีการแปลงค่าไปเป็น dictionary อีกครั้งด้วยคำสั่ง **dict\(\)** ดังตัวอย่างด้านล่าง


```python
a1 = dict(z1)
print (a1)
```

```python
dict(z1)
```

ข้อมูลจากคำสั่ง zip เมื่อแปลงข้อมูลไปแล้ว ข้อมูลดังกล่าวจะหายไป

```python

```

### Contruct Dictionary from 2 Dimension
Need 2 Dimension Data for construct Dictionary

```python
a2 = dict([['a', 12],['b', 34]])
```

```python
a2
```

**create dict  as  d1**

```python

```

```python

```

### Method ของ Dictionary



**clear\( \)** เป็นฟังก์ชันที่ใช้ในการลบค่าข้อมูลใน memory


```python
a1.clear()
print(a1)
```

```python

```

```python

```

```python

```

Dictionary  สามารถสร้างด้วยคำสั่งเป็น loops ได้  



```python
for i in range(len(names)):
    a1[names[i]] = numbers[i]
print(a1)
```

```python

```

```python

```

```python

```

**values\( \)** เป็นฟังก์ชันที่จะคืนค่า value ที่เก็บอยู่ใน dictionary ทั้งหมด


```python
a1.values()
```

```python

```

```python

```

```python

```

**keys\( \)** เป็นฟังก์ชันที่ค่า key หรือ index ที่ใช้ในการเข้าถึงค่าใน dictionary ทั้งหมด


```python
a1.keys()
```

```python

```

```python

```

```python

```

**items\( \)** เป็นฟังก์ชัที่คืนค่า ทั้ง keys และ values พร้อมกันอยู่ภายใต้ tuple  โดยผลลัพทธ์ที่ได้นี้จะเหมือนกับการใช้คำสั่ง zip


```python
a1.items()
```

```python

```

```python

```

```python

```

**pop\( \)** function เป็นฟังก์ชันที่จะใช้ในการเอาสมาชิกที่ต้องการออกจาก dictionary และสามารถส่งค่าไปยังตัวแปรใหม่ได้ โดยจะถูกส่งค่าไปเฉพาะส่วนที่เป็น value เท่านั้น ส่วน key จะละทิ้งไป

ดังตัวอย่างด้านล่างนี้



```python
a2 = a1.pop('Four')
print(a1)
print(a2)
```

```python

```

```python

```

## Exercise



### Ex 1

จงเขียนโปรแกรมเพื่อแปลงข้อมูล list 2 ตัวที่มีขนาดเท่ากัน ไปเป็น dictionary

เช่น

keys = \['สิบ', 'ยี่สิบ', 'สามสิบ'\]

values = \[10, 20, 30\]

คำตอบที่คาดหวัง

{'สิบ':10, 'ยี่สิบ':20, 'สามสิบ':30}



<!-- grade_id: 12d7b2, points: 0, grade: false, solution: true -->
```python
def exercise1(keys = ['สิบ', 'ยี่สิบ', 'สามสิบ'], values = [10, 20, 30]):
    myDict = None # Your code here
    ### your code here

    print(myDict)
    return myDict
```

<!-- grade_id: a59f38, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:

assert exercise1() == {'สิบ': 10, 'ยี่สิบ': 20, 'สามสิบ': 30}
assert exercise1(['สิบ','ยี่สิบ'],[10,20])== {'สิบ': 10, 'ยี่สิบ': 20}

```

### Ex 2

หากต้องการสร้างเครื่องชั่งน้ำหนักผลไม้ที่สามารถคำนวณราคาออกมาได้

จงเขียนโปรแกรมเพื่อจัดเก็บราคาผลไม้ต่อกิโลกรัมในรูปแบบของ dictionary  ดังต่อไปนี้

กล้วยราคา 20 บาท/kg

แตงโมงราคา 30 บาท/kg

ส้มราคา 40 บาท/kg

โดยโปรแกรมจะรับค่าน้ำหนักสินค้าที่ซื้อและแสดงออกมาเป็นราคาสินค้าที่ต้องชำระ



<!-- grade_id: f9bf9e, points: 0, grade: false, solution: true -->
```python
def exercise2(weight = 2, fruit = 'banana'):
    price ={'banana':20, 'watermelon':30, 'orange':40 }
    totalPrice= None # your code here
    ### your code here


    print("Total price %s = %0.2f" %(fruit.capitalize(),totalPrice))
    return totalPrice
```

<!-- grade_id: 36af12, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
assert exercise2() == 40
assert exercise2(2,'orange') == 80
```

```python

```