<a href="https://colab.research.google.com/github/toche7/BasicPython/blob/main/08_ImportandLib.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 08: Important Python Libraries

Revised: T. Samanchuen <br>
Version: 004



# Numpy

## การสร้าง Array

ในตัวอย่างข้างล่าง เราจะทดลองสร้าง Numpy array ขึ้นมาอย่างง่ายๆ จาก Python list



```python
import numpy as np        # วิธีการที่จะเรียกใช้ numpy library โดยกำหนดชื่อย่อว่า np
```

```python
a = np.array([1, 2, 3])   # สร้าง array 1 มิติ (เวกเตอร์) จาก Python list [1, 2, 3]

```

```python
print(type(a))            # พิมพ์ "<class 'numpy.ndarray'>"
```

```python
print(a.shape)
```

```python
print(a[0], a[1], a[2])   # พิมพ์ "1 2 3"

```

```python
a[0] = 5                  # เปลี่ยนค่าใน array
print(a)                  # พิมพ์ "[5, 2, 3]"

```

```python
b = np.array([[1,2,3],[4,5,6]])    # สร้าง array 2 มิติ (เมตริกซ์)
print(b)

```

```python
c = np.array([[[1],[2],[3]],[[4],[5],[6]]])
```

```python
c.shape
```

```python
print(b.shape)                     # พิมพ์ "(2, 3)" แปลว่าเป็น array 2 มิติ ที่มีความยาวเท่ากับ 2 และ 3 ตามลำดับ
                                   # พูดง่ายๆ ว่าเป็นเมตริกซ์ขนาด 2x3 นั่นเอง
print(b[0, 0], b[0, 1], b[1, 0])   # พิมพ์ "1 2 4"
```

```python
d = np.arange(10)
```

```python
d
```

```python
print(type(d))
```

```python
a = np.zeros((2,2))   # สร้าง array 2 มิติขนาด 2x2 ที่มีสมาชิกเป็น 0 ทั้งหมด
print(a)              # พิมพ์ "[[ 0.  0.]
                      #        [ 0.  0.]]"
```

```python
b = np.ones((1,2))    # สร้าง array 2 มิติขนาด 1x2 ที่มีสมาชิกเป็น 1 ทั้งหมด
print(b)              # พิมพ์ "[[ 1.  1.]]"


```

```python
c = np.full((2,2), 7)  # สร้าง array 2 มิติขนาด 2x2 ที่มีสมาชิกเป็นค่าคงที่เท่ากับ 7 ทั้งหมด
print(c)               # พิมพ์ "[[ 7.  7.]
                       #       [ 7.  7.]]"
```

```python
e = np.random.random((2,2))  # สร้าง array 2 มิติขนาด 2x2 ที่มีสมาชิกเป็นค่าสุ่ม (random numbers) ในช่วง (0,1)
print(e)                     # พิมพ์ค่าสุ่มอาทิเช่น "[[ 0.91940167  0.08143941]
                             #               [ 0.68744134  0.87236687]]"
```

## โอเปอร์เรชั่นทางคณิตศาสตร์ของ Array

### Element-wise operations
### การบวก ลบ คูณ หาร  "รายตัว" กรณี Array ทั้งสองมีขนาดเท่ากัน
การบวก ลบ คูณ หารในสมาชิกแต่ละตัวของ array สอง array นั่นคือบวก ลบ คูณ หารเฉพาะสมาชิกในตำแหน่งเดียวกัน สามารถทำได้อย่างตรงไปตรงมาในกรณีที่ array ทั้งสองมีขนาดเท่ากัน

```python
x = np.array([[1,2],[3,4]])
print(x)
```

```python
y = np.array([[5,6],[7,8]])
print(y)
```

```python
print(x + y)
```

```python
print(x - y)
```

```python
print(x * y)
```

ตัวอย่างการหาร เมื่อเลขไม่ลงตัว สมาชิกใน array จะเปลี่ยนเป็นจำนวนจริงโดยอัตโนมัติเหมือนกรณีการหารตัวเลข (ดูหัวข้อบน)

```python
print(x / y)

 [[ 0.2         0.33333333]
 [ 0.42857143  0.5       ]]
```



```python
print(x / y)
```

การยกกำลัง และถอดรากในสมาชิกแต่ละตัว ก็ทำได้ตรงไปตรงมาดังตัวอย่าง

```python
print(x ** 2)

```

```python
r = np.array([[1,2],[3,4]])
print(r ** 3)

```

```python
np.power(r,3)
```

ถอดรากที่สองมีสองวิธีคือเรียกฟังก์ชั่นของ numpy หรือ ยกกำลังด้วย 1/2

```python
print(np.sqrt(x))
```

```python
print(x ** (1/2))
```

#### EX 1

หากมีข้อมูลนำเข้าเป็น Array ของรัศมีของลูกโป่งที่บรรจุแก๊ส เช่น $r$ = [2, 3, 4, 5] จงเขียนโปรแกรมเพื่อคำนวณหาปริมาณแก๊ส ทั้งหมดของลูกโป่ง
สมมุติว่าลูกโป่งมีลักษณะเป็นทรงกลม โดยปริมาตรของทรงกลมมีสูตรเป็น
$$
V = \frac{4}{3}\pi r^3
$$
Hint: <br>
 $\pi$ ใน Numpy ใช้เป็น np.pi



<!-- grade_id: a6a626, points: 0, grade: false, solution: true -->
```python
def Spherical(r):  # modify function name and parameters
    """
    Compute the sum of the Volumn of Spherical.
    r : array of radius
    """
    ### BEGIN SOLUTION
    volumn = (r ** 3)*np.pi*4/3
    ### END SOLUTION
    return [np.sum(volumn)]
```

<!-- grade_id: e14f8b, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(Spherical(np.array([1])), [4.1887902047863905])
assert_equal(Spherical(np.array([2,3])), [146.60765716752368])
```

```python
r = np.array([[1,2],[3,4]])
Spherical(r)
```

## Array indexing
Numpy มีวิธี index array ได้หลายรูปแบบ

### Slicing
เราสามารถใช้ slicing ได้เช่นเดียวกับ Python list และเลือกใช้บนมิติใดของ array ก็ได้

เช่น สร้าง array 2มิติ ที่มีขนาด (3, 4)

```python
a = np.array([[1,2,3,4], [5,6,7,8], [9,10,11,12]])
print(a)
```

ใช้ slicing เพื่อดึงสองแถวแรก และคอลัมภ์ที่สองและสาม สังเกตว่า index เริ่มจาก 0 ดังนั้นคอลัมภ์ที่สองและสามจึงมี index 1 และ 2 และอย่าลืมว่าใน syntax ของ Python จะนำ index ตัวสุดท้ายไปลบด้วยหนึ่ง

```python
b = a[:2, 1:3]
print(b)
```

```python
s = a.shape
print(s)
```

```python
print(s[1])
```

```python
print(s[0])
```

นอกจากนี้ b ยังเสมือนเป็น pointer ที่ชี้ไปยังข้อมูลชุดเดียวกับ a ดังนั้นถ้าเปลี่ยนค่าใน b ซึ่งเป็น slice ของ a ข้อมูลใน a ก็จะเปลี่ยนด้วย

```python
print(a[0, 1])   # พิมพ์ "2"
b[0, 0] = 77     # b[0, 0] is the same piece of data as a[0, 1]
print(a[0, 1])   # พิมพ์ "77"
```

```python
a = np.array([[1,2,3,4], [5,6,7,8], [9,10,11,12]])
print(a)
```

สังเกตว่าการ index แถวที่สองทำได้สองวิธีคือ ใช้ index 1 หรือใช้ slicing 1:2 ซึ่งใน syntax ของ python จะให้ผลลัพธ์ที่ **แตกต่างกัน**

การใช้ index เป็นเพียงตัวเลข เช่น 1 จะลดมิติของ array ลง 1 มิติ (ดูตัวอย่างข้างล่าง)
ในขณะที่การใช้ slicing index เช่น 1:2 เราจะได้ผลลัพธ์เป็น array ที่มีมิติเท่าเดิม

```python
row_r1 = a[1, :]    # ผลลัพธ์คือ array 1 มิติ
row_r2 = a[1:2, :]  # ผลลัพธ์คือ array 2 มิติ
print(row_r1, row_r1.shape)  # พิมพ์ "[5 6 7 8] (4,)"
print(row_r2, row_r2.shape)  #พิมพ์ "[[5 6 7 8]] (1, 4)"

```

#### EX 2

หากมีข้อมูลที่เป็น array 2 มิติ จงเขียนโปรแกรมที่สามารถตัดเอาเฉพาะข้อมูล 2 Column สุดท้ายของ Array เพื่อนำไปใช้ต่อ



<!-- grade_id: ab1671, points: 0, grade: false, solution: true -->
```python
def slicerLastTwoColumn(a):  # modify function name and parameters
    """
    Compute the squares of the numbers from 1 to n.  [replace with function description]
    """
    s = a.shape
    if s[1] < 3: raise ValueError("n must be at least 3")
    ### BEGIN SOLUTION
    #raise NotImplementedError()
    b = a[:, (s[1] - 2): s[1]]
    ### END SOLUTION
    return b
```

<!-- grade_id: 32ca09, points: 1, grade: true, solution: false -->
```python
a = np.array([[1,2,3,4], [5,6,7,8], [9,10,11,12]])
b = np.array([[1,2,3], [5,6,7], [9,10,11]])

# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
# from nose.tools import assert_array_equal
np.testing.assert_array_equal(slicerLastTwoColumn(a), np.array([[3,  4],[7,  8],[11, 12]]))
np.testing.assert_array_equal(slicerLastTwoColumn(b), np.array([[2,  3],[6,  7],[10, 11]]))
```

# Matplotlib
Matplotlib เป็นไลบรารีสำหรับการพล็อดและจัดการรูปภาพ ซึ่งมีการใช้งานคล้ายกับบน MATLAB.


## Plotting
ตัวอย่งการพล็อดกราฟสองมิติ

```python
import numpy as np
import matplotlib.pyplot as plt

```

```python
# สร้าง x ให้เป็น array ของตัวเลขตั้งแต่ 0 ไปจนถึง 3pi โดยแต่ละจุดห่างกันเท่ากับ 0.1
x = np.arange(0, 6 * np.pi, 0.1)
# กำหนด y ให้เป็น array มีค่าเท่ากับไซน์ของ x ในแต่ละจุดซึ่งเรียกโดยใช้ฟังก์ชัน np.sin
y = np.sin(x)

```

```python
# ระบุว่าเราต้องการพล็อดกราฟของ x และ y
plt.plot(x, y)
plt.show()  # เมื่อระบุสิ่งที่เราต้องการพล็อตครบแล้ว สามารถแสดงผลการพล็อตได้ด้วยคำสั่ง plt.show()
```

จากโค้ดในเซลล์ที่แล้ว เราสามารถเพิ่มชื่อกราฟ, คำอธิบาย และชื่อแกนได้ง่ายๆ

```python
import numpy as np
import matplotlib.pyplot as plt

# สร้างกราฟทั้งไซน์และโคไซน์ของ x
x = np.arange(0, 3 * np.pi, 0.1)
y_sin = np.sin(x)
y_cos = np.cos(x)

# สั่งให้พล็อตทั้งสองกราฟพร้อมกัน
plt.plot(x, y_sin)
plt.plot(x, y_cos)
plt.xlabel('x axis label') # กำหนดชื่อบนแกน x
plt.ylabel('y axis label') # กำหนดชื่อบนแกน y
plt.title('Sine and Cosine') # กำหนดชื่อไตเติ้ล
plt.legend(['Sine', 'Cosine']) # กำหนดลาเบลของกราฟแต่ละเส้น (ตามลำดับคำสั่ง plt.plot)
plt.show() # เมื่อระบุสิ่งที่เราต้องการพล็อตครบแล้ว สามารถแสดงผลการพล็อตได้ด้วยคำสั่ง plt.show()
```

## Subplot
เราสามารถแบ่งรูปภาพเป็นหลายส่วน และพล็อตแต่ละกราฟในแต่ละส่วนย่อยได้ด้วยฟังก์ชัน `supplot`

https://matplotlib.org/stable/gallery/subplots_axes_and_figures/subplots_demo.html

การใช้คำสั่ง Subplot สามารถแบ่งได้สองลักษณะใหญ่ๆ คือ ส่ง plot ผ่านคำสั่ง class plt  และอีกวิธีหนึ่งก็คือสร้างเป็น object ขึ้นมาก่อนแล้วจึงจะสั่งผ่าน object อีกที



### สั่ง plot ผ่าน object ax
จริงๆ สามารถตั้งชื่อ  object เป็นอย่างอื่นได้ แต่เพื่อเป็นมาตรฐานจึงจะใช้เช่นเดียวกับที่ใช้กันโดยทั่วไป มีวิธีการดำเนินการดังนี้
1. ประกาศ object  fig, ax = plot.subplot(x,y)  โดย x และ  y  เป็นขนาด เมตริก  ของ subplot ถ้าไม่ระบุก็จะมีขนาดเป็น 1
2. สั่ง plot ผ่าน  ax.plot()
3. ตั้งชื่อผ่าน ax.set_title()


```python
import numpy as np
import matplotlib.pyplot as plt

# เตรียมข้อมูลเหมือนในตัวอย่างที่แล้ว
x = np.arange(0, 3 * np.pi, 0.1)
y_sin = np.sin(x)
y_cos = np.cos(x)

fig, ax = plt.subplots()  # fig จะเป็น object แบบ global และ ax จะเป็น object ในแนวแกน
ax.plot(x, y_sin)
ax.set_title('A single plot')
```

```python
import matplotlib.pyplot as plt
import numpy as np

# Create a Figure with 2 rows and 2 columns of subplots:
fig, ax = plt.subplots(2, 2, constrained_layout = True)
#fig.tight_layout(pad=2.0)

x = np.linspace(0, 5, 100)

# Index 4 Axes arrays in 4 subplots within 1 Figure:

ax[0, 0].plot(x, np.sin(x), 'g') #row=0, column=0
ax[0, 0].set_xlabel('x1')
ax[0, 0].set_ylabel('y1')
ax[0, 0].set_title('Ax[0,0]')

ax[1, 0].plot(range(100), 'b') #row=1, column=0
ax[1, 0].set_xlabel('x1')
ax[1, 0].set_ylabel('y1')
ax[1, 0].set_title('Ax[1,0]')

ax[0, 1].plot(x, np.cos(x), 'r') #row=0, column=1
ax[0, 1].set_xlabel('x1')
ax[0, 1].set_ylabel('y1')
ax[0, 1].set_title('Ax[0,1]')

ax[1, 1].plot(x, np.tan(x), 'k') #row=1, column=1
ax[1, 1].set_xlabel('x1')
ax[1, 1].set_ylabel('y1')
ax[1, 1].set_title('Ax[1,1]')



plt.show()
```

#### EX 3

จงวาดกราฟของ ฟังก์ชัน $ y = x^2 + 2x + 1$



<!-- grade_id: 26d95c, points: 1, grade: true, solution: true -->
```python
##  YOUR ANSWER HERE

# YOUR CODE HERE
### BEGIN SOLUTION
import numpy as np
import matplotlib.pyplot as plt

fig, ax = plt.subplots(constrained_layout=True)
# เตรียมข้อมูลเหมือนในตัวอย่างที่แล้ว
x = np.arange(-5, 5, 0.1)
y = x**2 + 2*x + 1
ax.plot(x, y)
ax.set_xlabel('x')
ax.set_ylabel('y')
ax.set_title('My Function')
plt.legend(['$X^2$'])
plt.show()
### END SOLUTION

```